The magnitude and time course of pre-saccadic foveal prediction depend on the conspicuity of the saccade target

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

This study examines to what extent this phenomenon varies based on the visibility of the saccade target. Visibility is defined as the contrast level of the target with respect to the noise background, and it is related to the signal-to-noise ratio of the target. A more visible target facilitates the oculomotor behavior planning and execution, however, as speculated by the authors, it can also benefit foveal prediction even if the foveal stimulus visibility is maintained constant. Remarkably, the authors show that presenting a highly visible saccade target is beneficial for foveal vision as detection of stimuli with an orientation similar to that of the saccade target is improved, the lower is the saccade target visibility, the less prominent is this effect.

Strengths:

The results are convincing and the research methodology is technically sound.

Weaknesses:

It is still unclear why the pre-saccadic enhancement would oscillate for targets with higher opacity levels, and what would be the benefit of this oscillatory pattern. The authors do not speculate too much on this and loosely relate it to feedback processes, which are characterized by neural oscillations in a similar range.

We thank the reviewer for their assessment. We intentionally decided to describe the oscillatory pattern without claiming to be able to pinpoint its origin. The finding was incidental and, based on psychophysical data alone, we would not feel comfortable doing anything but loosely relating it to potential mechanisms on an explicitly speculative basis. In the potential explanation we provide in the manuscript, the oscillatory pattern would likely not serve a benefit–rather, it would constitute an innate consequence and, thus, a coincidental perceptual signature of potential feedback processes.

Reviewer #2 (Public review):

Summary:

In this manuscript, the authors ran a dual task. Subjects monitored a peripheral location for a target onset (to generate a saccade to), and they also monitored a foveal location for a foveal probe. The foveal probe could be congruent or incongruent with the orientation of the peripheral target. In this study, the authors manipulated the conspicuity of the peripheral target, and they saw changes in performance in the foveal task. However, the changes were somewhat counterintuitive.

We regret that our findings remain counterintuitive to the reviewer even after our extensive explanations in the previous revision round and the corresponding changes in the manuscript. We repeat that both the decrease in foveal Hit Rates and the increase in foveal enhancement with increasing target contrast were expected and preregistered prior to data collection.

Strengths:

The authors use solid analysis methods and careful experimental design.

Comments on revisions:

The authors have addressed my previous comments.

One minor thing is that I am confused by their assertion that there was no smoothing in the manuscript (other than the newly added time course analysis). Figure 3A and Figure 6 seem to have smoothing to me.

When the reviewer suggested that the “data appear too excessively smoothed” in the first revision, we assumed that they were referring to pre-saccadic foveal Hit and False Alarm rates, not to fitted distributions. As we state in the legend of Figure 3A (as well as in Figures 6 and S1), the “smoothed” curves constitute the probability density distributions of our raw data. Concerning the energy maps resulting from reverse correlation analyses, we described our proceeding in detail in our initial article (Kroell & Rolfs, 2022):

“Using this method, we obtained filter responses for 260 SF*ori combinations per noise image (Figure 6 in Materials and methods, ‘Stimulus analysis’). SFs ranged from 0.33 to 1.39 cpd (in 20 equal increments). Orientations ranged from –90–90° (in 13 equal increments). To normalize the resulting energy maps, we z-transformed filter responses using the mean and standard deviation of filter responses from the set of images presented in a certain session. To obtain more fine-grained maps, we applied 2D linear interpolations by iteratively halving the interval between adjacent values 4 times in each dimension. To facilitate interpretability, we flipped the energy maps of trials in which the target was oriented to the left. In all analyses and plots,+45° thus corresponds to the target’s orientation while –45° corresponds to the other potential probe orientation. Filter responses for all response types are provided at https://osf.io/v9gsq/.”

We have added a pointer to this explanation to the current manuscript (see line 836).

Another minor comment is related to the comment of Reviewer 1 about oscillations. Another possible reason for what looks like oscillations is saccadic inhibition. when the foveal probe appears, it can reset the saccade generation process. when aligned to saccade onset, this appears like a characteristic change in different parameters that is time-locked to saccade onset (about a 100 ms earlier). So, maybe the apparent oscillation is a manifestation of such resetting and it's not really an oscillation. so, I agree with Reviewer 1 about removing the oscillation sentence from the abstract.

While we understand that a visible probe will result in saccadic inhibition (White & Rolfs, 2016), we are unsure how a resetting of the saccade generation process should manifest in increased perceptual enhancement of a specific, peripheral target orientation in the presaccadic fovea. Moreover, as we describe in our initial article (Kroell & Rolfs, 2022), we updated the background noise image every 50 ms and embedded our probe stimulus into the surrounding noise using smooth orientation filters and raised cosine masks to avoid a disruptive influence of probe appearance on movement planning and execution (Hanning, Deubel, & Szinte, 2019). And indeed, we demonstrated that the appearance of the foveal probe did not disrupt saccade preparation, that is, did not increase saccade latencies compared to ‘probe absent’ trials in which no foveal probe was presented (see Kroell & Rolfs, 2022; sections “Parameters of included saccades in Experiment 1” and “Parameters of included saccades in Experiment 2”). In the current submission, saccade latencies in ‘probe present’ trials exceeded saccade latencies in ‘probe absent’ trials by a mere 4.7±2.3 ms. Additionally, to inspect the variation of saccade execution frequency directly, we aligned the number of saccade generation instances to the onset of the foveal probe stimulus (see Author response image 1). In line with what we described in a previous paradigm employing flickering bandpass filtered noise patches (Kroell & Rolfs, 2021; 10.1016/j.cortex.2021.02.021), we observed a regular variation in saccade execution frequency that reflected the duration of an individual background noise image (50 ms in this investigation). In other words, the repeated dips in saccadic frequency are likely caused by the flickering background noise and not the onset of the foveal probe which would produce a single dip ~100 ms after probe onset. Given these results, we do not see a straight-forward explanation for how the variation of saccade execution frequency in 20 Hz intervals would boost peripheral-to-foveal feature prediction before the saccade in ~10 Hz intervals. Nonetheless, we removed the sentence referencing oscillations from the Abstract.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Overall, The authors did a good job in addressing the points I raised. Two new sections were added to the manuscript, one to address how the mechanisms of foveal predictions would play out in natural viewing conditions, and another one examining more in depth the potential neural mechanisms implicated in foveal predictions. I found these two sections to be quite speculative, and at points, a bit convoluted but could help the reader get the bigger picture. I still do not have a clear sense of why the pre-saccadic enhancement would oscillate for targets with higher opacity levels, and what would be the benefit of this oscillatory pattern. The authors do not speculate too much on this and loosely relate it to feedback processes, which are characterized by neural oscillations in a similar range.

Please see our response to ‘Weaknesses’.

I still find this a loose connection and would suggest removing the following phrase from the abstract "Interestingly, the temporal frequency of these oscillations corresponded to the frequency range typically associated with neural feedback signaling".

We have removed this phrase.

Finally, the authors should specify how much of this oscillation is due to oscillations in HR of cong vs. oscillations in HR of incongruent trials or both.

We fitted separate polynomials to congruent and incongruent Hit Rates instead of their difference. Peaks in enhancement relied on both, oscillatory increases in congruent Hit Rates and simultaneous decreases in incongruent Hit Rates. In other words, enhancement peaks appear to reflect a foveal enhancement of target-congruent feature information along with a concurrent suppression of target-incongruent features. We added this paragraph and Figure 4 to the Results section.

Additional changes:

Two figures had accidentally been labeled as Figure 5 in our first revision. We corrected the figure legends and all corresponding figure references in the text.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: We previously demonstrated that during the preparation of a large-scale saccadic eye movement, defining features of the eye movement target (i.e., its orientation) are anticipated in pre-saccadic foveal vision (Kroell & Rolfs, 2022). In this Research Advance, we show that the conspicuity of orientation information at the saccade target location influences the magnitude and time course of foveal enhancement. As in our previous study, observers prepared a saccade to a peripheral orientation signal (the target) while monitoring the appearance of a second orientation signal (the probe) in their pre-saccadic center of gaze. The probe appeared in 50% of trials and either had the same orientation as the target (congruent) or a different orientation (incongruent). In the current study, we manipulated the opacity of the target against the 1/f background noise in four logarithmic steps from 25% to 90%. An increase in opacity translated to an increase in luminance contrast and the signal-to-noise ratio of orientation information within the target region. We made three main observations: First, foveal Hit Rates for target-congruent and incongruent probes decreased as target opacity increased, presumably since attention was increasingly drawn to the target the more salient it became. Crucially, foveal enhancement defined as the difference between congruent and incongruent Hit Rates increased with opacity. Second, the time course of foveal enhancement appeared to show an oscillatory pattern that was particularly pronounced at higher target opacities. Interestingly, the temporal frequency of these oscillations corresponded to the frequency range typically associated with neural feedback signaling. Third, unlike the difference in Hit Rates, the difference in False Alarm Rates did not vary with opacity. Instead, reverse correlations suggest that at higher target opacities, False Alarms were increasingly triggered by signal, that is, by incidental orientation information in the foveal noise. Beyond providing new mechanistic insights into active foveal processing, these findings are relevant for researchers planning to adapt our paradigm to study related questions. Presenting the saccade target at a high signal-to-noise ratio appears beneficial as congruency effects, especially when time-resolved, are most robustly detectable.

Revised version: We previously demonstrated that during the preparation of a large-scale saccadic eye movement, defining features of the eye movement target (i.e., its orientation) are anticipated in pre-saccadic foveal vision (Kroell & Rolfs, 2022). In this Research Advance, we show that the conspicuity of orientation information at the saccade target location influences the magnitude and time course of foveal enhancement. As in our previous study, observers prepared a saccade to a peripheral orientation signal (the target) while monitoring the appearance of a second orientation signal (the probe) in their pre-saccadic center of gaze. The probe appeared in 50% of trials and either had the same orientation as the target (congruent) or a different orientation (incongruent). In the current study, we manipulated the opacity of the target against the 1/f background noise in four logarithmic steps from 25% to 90%. An increase in opacity translated to an increase in luminance contrast and the signal-to-noise ratio of orientation information within the target region. We made three main observations: First, foveal Hit Rates for target-congruent and incongruent probes decreased as target opacity increased, presumably since attention was increasingly drawn to the target the more salient it became. Crucially, foveal enhancement defined as the difference between congruent and incongruent Hit Rates increased with opacity. Second, the time course of foveal enhancement showed an oscillatory pattern that was particularly pronounced at higher target opacities. Third, unlike the difference in Hit Rates, the difference in False Alarm Rates did not vary with opacity. Instead, reverse correlations suggest that at higher target opacities, False Alarms were increasingly triggered by signal, that is, by incidental orientation information in the foveal noise. Beyond providing new mechanistic insights into active foveal processing, these findings are relevant for researchers planning to adapt our paradigm to study related questions. Presenting the saccade target at a high signal-to-noise ratio appears beneficial as congruency effects, especially when time-resolved, are most robustly detectable.

Introduction

We previously demonstrated that during the preparation of a saccadic eye movement to a peripheral visual field location, defining features of the eye movement target are predictively enhanced in pre-saccadic foveal vision (Kroell & Rolfs, 2022). In this Research Advance, we manipulated the conspicuity of the saccade target in a parametric fashion and, based on a large number of trials, describe the influence of this manipulation on the magnitude and temporal development of foveal prediction.

In our original paradigm, we presented a rapid sequence of 1/f noise images filling the entire screen (see Hanning & Deubel, 2022a; 2022b). Human observers maintained fixation in the screen center until an orientation signal, generated by filtering the luminance content at the desired display location to either a −45° or 45° tilt, appeared 10 degrees of visual angle (dva) to the left or the right of the screen center. Observers executed an eye movement to the peripheral orientation signal (i.e., the target) as soon as they detected it. During saccade preparation, a second orientation signal (the foveal probe) could appear in observers’ current, pre-saccadic center of gaze on 50% of trials. If presented, the probe either had the same orientation as the target (congruent) or a different orientation (incongruent). After executing the saccade, observers reported if they had perceived the probe in the screen center before moving their eyes and, if so, which orientation they had perceived.

Earlier version: Starting 175 ms before saccade onset, observers’ Hit Rates (HRs) for foveal probes with target-congruent orientation significantly exceeded incongruent HRs. This congruency effect was spatially confined to the center of gaze and its immediate vicinity and more pronounced during saccade preparation than during passive fixation. Based on these findings, we proposed that feedback connections from higher-order visual areas relay relevant saccade target features to foveal retinotopic cortex (see also Williams et al., 2008; Chambers, Allen, Maizey, & Williams, 2013; Fan, Wang, Shao, Kersten, & He, 2016; Weldon, Rich, Woolgar, & Williams, 2016; Yu & Shim, 2016; Weldon, Woolgar, Rich & Williams, 2020). Fed-back orientation signals subsequently combine with feedforward foveal input to early visual cortex and facilitate the detection of target-congruent feature information in the center of gaze.

Revised version: Starting 175 ms before saccade onset, observers’ Hit Rates (HRs) for foveal probes with target-congruent orientation significantly exceeded incongruent HRs. This congruency effect was spatially confined to the center of gaze and its immediate vicinity and more pronounced during saccade preparation than during passive fixation. Based on these findings, we proposed that feedback connections from higher-order visual areas enhance relevant saccade target features in foveal retinotopic cortex (see also Williams et al., 2008; Chambers, Allen, Maizey, & Williams, 2013; Fan, Wang, Shao, Kersten, & He, 2016; Weldon, Rich, Woolgar, & Williams, 2016; Yu & Shim, 2016; Weldon, Woolgar, Rich & Williams, 2020). Internal orientation predictions subsequently combine with feedforward foveal input to early visual cortex and facilitate the detection of target-congruent feature information in the center of gaze.

In all previous experiments, we presented the saccade target at an opacity of 60% against the background noise. Importantly, the choice of a certain target opacity in our paradigm dictates the luminance contrast as well as the signal-to-noise ratio (SNR) of orientation information within the saccade target region. We had chosen a medium target opacity in previous experiments since we presumed that feature information may need to be relevant for the movement task to be predicted foveally. In other words, while high target opacities would likely have generated a salient and immediate target pop-out, continuous peripheral sampling of orientation information (see Ludwig, Davies, & Eckstein, 2013) may have been necessary to program the eye movement at a medium opacity. On the other hand, a highly conspicuous saccade target may not only facilitate the oculomotor task and free up resources for perceptual detection but could benefit foveal prediction more directly: The higher the SNR of peripheral orientation information, the less noisy the foveally predicted signal may be. Moreover, the timing of foveal feedback effects during fixation varied considerably across previously published investigations (from 117 ms in Weldon et al. (2016; 2020) to 550 ms in Fan et al. (2016)). While this variation has been ascribed to a flexible adjustment of the feedback mechanism to current task demands (Fan et al., 2016), low-level stimulus features may play an equally important role, at least in our investigations. In particular, neuronal response latencies decrease systematically as the contrast of visual input increases. While this phenomenon is reliably observed at varying stages of the visual processing hierarchy—such as the lateral geniculate nucleus (Lee, Elepfandt, & Virsu, 1981b), primary visual cortex (e.g., Albrecht, 1995; Carandini & Heeger, 1994; Carandini, Heeger, & Movshon, 1997; Carandini, Heeger, & Senn, 2002), and anterior superior temporal sulcus (STSa; Oram, Xiao, Dritschel, & Payne, 2002; van Rossum, van der Meer, Xiao, & Oram, 2008)—influences of contrast on neuronal response latency are particularly pronounced in higher-order visual areas: A doubling of stimulus contrast has been shown to decrease the latency of V1 neurons by 8 ms, compared to a reduction of 33 ms in area STSa (Oram et al., 2002; van Rossum et al., 2008). Assuming that the peripheral target is processed in a bottom-up fashion until it reaches higher-order object processing areas, the time point at which peripheral signals are available for feedback should be dictated by the temporal dynamics of visual feedforward processing.

In this Research Advance, we describe the influence of saccade target opacity (Throughout the paper, we use the term ‘opacity’ when we refer to the experimental manipulation (that is, a variation of the transparency, i.e., 1-opacity of the target patch against the background noise) and the term ‘conspicuity’ when we discuss our findings conceptually.) on pre-saccadic foveal congruency effects. For this purpose, we varied target opacity in four steps (equally spaced on a log scale: 25.0%, 38.3%, 58.7%, 90.0%) while leaving the remaining task parameters unchanged (Figure 1). An increase in target opacity from 25% to 90% corresponded to an increase in Michelson contrast (Michelson, 1927) from .53 to .78 and an increase in the SNR of the filtered orientation from 1.40 to 2.58 (see Methods). Note that the foveal probe was presented at the same opacity (individually adjusted for every session and observer) across all target opacity conditions. The difference between congruent and incongruent HRs for foveal probes, i.e., enhancement, increased monotonically with target opacity. This finding endorses the assumption that predictive foveal processing can support perceptual continuity in natural visual environments, in which saccades are routinely directed towards the currently most conspicuous object in the scene (‘t Hart, Schmidt, Roth, & Einhäuser, 2013). Moreover, especially when systematic variations of saccade latency were taken into account, enhancement (Note that throughout the manuscript, the term ‘enhancement’ is used to refer to the difference between congruent and incongruent Hit Rates (HRcong – HRincong)) required less and less time to develop as the opacity of the saccade target, and thus its contrast, increased. These findings not only provide new mechanistic insights but refine our method and are relevant for researchers planning to use or adapt our paradigm to study related questions.

Removed in the revised version: Summary of the paradigm.

Earlier version: A. Example trial procedure: The saccade target and foveal probe were embedded in full-screen noise images flickering at a frequency of 20 Hz (image duration of 50 ms). After 200 ms, the saccade target (an orientation-filtered patch; filtered to either –45° or +45°; 3 dva in diameter) appeared 10 dva to the left or the right of the screen center, cueing the eye movement. On 50% of trials, a probe (a second orientation-filtered patch; filtered to either –45° or +45°) appeared in the screen center at an early (top panel; highlighted element with dark blue outline), medium (light blue outline) or late (green outline) stage of saccade preparation. The foveal probe was presented for 50 ms and could be oriented either congruently or incongruently to the target. In contrast to our previous investigation, the saccade target was presented at one of four opacities (from 25% to 90%; in different blocks). B. An increase in target opacity translates to an increase in signal-to-noise ratio (SNR; left panel). Within a single trial, the increase in SNR at the target location manifests from the fourth noise image on (i.e., after the target appears; right panel; error bands correspond to the standard deviation across all images). C. Probe and target timing. Probe timing (left panel): histogram of time intervals between probe offset and saccade onset. Bar heights and error bars indicate the mean and standard error of the mean (SEM; n=9) across observers, respectively. On included trials, the probe appeared after target onset and therefore during saccade preparation (‘sac prep’). Trials in which the probe disappeared more than 250 ms before saccade onset (light grey), during the saccade (yellow) or after saccade offset (orange) were excluded. The yellow background rectangle illustrates the median saccade duration. Target timing (right panel): histogram of time intervals between target offset and saccade offset. Bar heights and error bars indicate the mean and standard error of the mean (SEM; n=9) across observers, respectively. Unlike in the previous study, we removed the target upon saccade initiation on all trials.

Revised version: Summary of the paradigm. A. Example trial procedure: The saccade target and foveal probe were embedded in full-screen noise images flickering at a frequency of 20 Hz (image duration of 50 ms). After 200 ms, the saccade target (an orientation-filtered patch; filtered to either –45° or +45°; 3 dva in diameter) appeared 10 dva to the left or the right of the screen center, cueing the eye movement. On 50% of trials, a probe (a second orientation-filtered patch; filtered to either –45° or +45°) appeared in the screen center at an early (top panel; highlighted element with dark blue outline), medium (light blue outline) or late (green outline) stage of saccade preparation. The foveal probe was presented for 50 ms and could be oriented either congruently or incongruently to the target. In contrast to our previous investigation, the saccade target was presented at one of four opacities (from 25% to 90%; in different blocks). B. An increase in target opacity translates to an increase in signal-to-noise ratio (SNR; left panel). Within a single trial, the increase in SNR at the target location manifests from the fourth noise image on (i.e., after the target appears; right panel; error bands correspond to the standard deviation across all images). C. Probe and target timing. Probe timing (left panel): histogram of time intervals between probe offset and saccade onset. Bar heights and error bars indicate the mean and standard error of the mean (SEM; n=9) across observers, respectively. On included trials, the probe appeared after target onset and therefore during saccade preparation (‘sac prep’). Trials in which the probe disappeared more than 250 ms before saccade onset (light grey), during the saccade (yellow) or after saccade offset (orange) were excluded. The yellow background rectangle illustrates the median saccade duration. Target timing (right panel): histogram of time intervals between target offset and saccade offset. Bar heights and error bars indicate the mean and standard error of the mean (SEM; n=9) across observers, respectively. Unlike in the previous study, we removed the target upon saccade initiation on all trials.

Results

Hit and False Alarm rates as a function of target opacity

Both congruent and incongruent HRs decreased with increasing target opacity (Figure 2A, left panel), arguably because attentional resources were increasingly drawn to the target the more salient it became. This decrease in foveal HRs, however, was less pronounced for probes with target-congruent orientation [slope of the linear regression line: −0.46% (congruent) vs −1.73% (incongruent) per opacity step, p < .001]. In all four opacity conditions, observers detected foveal probes more readily if they exhibited the same orientation as the saccade target, resulting in significantly higher congruent as compared to incongruent HRs (HRcong-incong = 3.3%, 3.4%, 5.4% & 6.8%, ps < .021; Figure 2A, middle panel). To fully characterize the impact of target opacity on enhancement in HRs, we performed a linear mixed-effects model in which we described the variance in enhancement (HRcong–HRincong) across all time points with a fixed effect of target opacity and a random intercept for observer (see Methods). Note that this model outperformed the simplest model including a fixed effect of target opacity only (ΔBIC = 11.6) and a more complex one involving a random intercept and random slope for observer (ΔBIC = 4.4). The fixed effect of target opacity reached significance, t(34) = 2.341, p = 0.025. More specifically, as the opacity of the saccade target increased, the difference between HRs for target-congruent and target-incongruent foveal probes increased as well, resulting in a significantly positive slope of the fitted linear regression line (slope = 1.3% per opacity step, p < .001; Figure 2A, right panel).

Removed in the revised version: Influence of target opacity on Hits and FAs across all time points.

Earlier version: A. Influence of target opacity on congruent and incongruent HRs (purple and gray data points in the left panel) as well as their difference (orange data points, middle panel). Lines and error bands correspond to the fitted linear regression lines ±2 standard errors of the mean (SEM). The slopes of the fitted regression line per observer (small circles) and their mean and SEM (big circle and error bar) are plotted in the right panel. Asterisks denote statistically significant comparisons (p < .05; determined via bootstrapping; n = 9 observers). B. Influence of target opacity on congruent and incongruent FARs (left panel) as well as their difference (middle and right panel). All conventions are as in A. C. Mean difference in filter energies around the target and non-target orientation for the lower two target opacities (left column) and the higher two target opacities (right panel). Lines connect the values of individual observers (small circles) in both conditions. Large circles and error bars denote the mean and SEM, respectively. D. Pearson correlation between the normalized slope in HRs (A, right panel) and the normalized slope in FARs (B, right panel). Circles indicate individual observers.

Revised version: Influence of target opacity on Hits and FAs across all time points. A. Influence of target opacity on congruent and incongruent HRs (purple and gray data points in the left panel) as well as their difference (orange data points, middle panel). Lines and error bands correspond to the fitted linear regression lines ±2 standard errors of the mean (SEM). The slopes of the fitted regression line per observer (small circles) and their mean and SEM (big circle and error bar) are plotted in the right panel. Asterisks denote statistically significant comparisons (p < .05; determined via bootstrapping; n = 9 observers). B. Influence of target opacity on congruent and incongruent FARs (left panel) as well as their difference (middle and right panel). All conventions are as in A. C. Mean difference in filter energies around the target and non-target orientation for the lower two target opacities (left column) and the higher two target opacities (right panel). Lines connect the values of individual observers (small circles) in both conditions. Large circles and error bars denote the mean and SEM, respectively. D. Pearson correlation between the normalized slope in HRs (A, right panel) and the normalized slope in FARs (B, right panel). Circles indicate individual observers.

On probe-absent trials, observers reported perceiving the target orientation more often than the non-target orientation for the 25%, 59% and 90% opacity condition, resulting in significantly higher congruent as compared to incongruent False Alarm Rates (FARs; FARcong-incong = 6.5%, 6.2% & 4.6%, ps < .02; Figure 2B). We performed a linear mixed-effects model to describe the difference between congruent and incongruent FARs with a fixed effect of target opacity and a random intercept for observer. Unlike for HRs, the fixed effect of target opacity did not affect the differences in FARs, t(34) = −0.44, p = .661. The slope of the fitted linear regression line was statistically indistinguishable from zero (−0.2% per opacity step, p = .713). Inspecting congruent and incongruent FARs instead of their difference revealed that, just like HRs, they both decreased with increasing target opacity. This decrease, however, was comparable for congruent and incongruent FARs (slopes: −1.9% vs −1.6% per opacity step, p = .715).

In our paradigm, the entire screen was covered with pink noise. In consequence, the display region the probe could appear in was never void of signal but contained incidental orientation information even on probe-absent trials. Trials in which observers reported perceiving the probe when no stimulus was presented may therefore contain a mixture of two response types: (i) ‘true’ FAs, that is, unsystematic or biased reports and (ii) ‘true Hits’, that is, systematic reactions to incidental orientation information in the foveal noise region. As target opacity (and, presumably, the strength of the fed-back signal) increased, the proportion of ‘true Hits’ within the subset of congruent FAs may have increased while the proportion of the ‘true FAs’ decreased. The relative contribution of both response types may thus have varied considerably without affecting the mean congruent FAR. Indeed, two observations suggest that the contribution of unsystematic responses to FAs in general, and congruent FAs in particular, decreased as target opacity increased.

First, we described the visual properties of the foveal noise on FA trials by convolving a set of Gabor filters with different spatial frequencies (SFs) and orientations with the pixel content in a 3 dva region around the pre-saccadic center of gaze. We then subtracted filter energies on incongruent FA trials from filter energies on congruent FA trials (see Methods). If FAs relied partially on signal, this difference image should be characterized by high filter energies around the target orientation and low filter energies around the non-target orientation. Interestingly, the difference in filter energies around the target and non-target orientation was significantly more pronounced for the higher as compared to the lower two target opacities (p = .008; Figure 2C). This observation suggests that the contribution of ‘true Hits’ to the pool of foveal FAs increased with target opacity, allowing systematic orientation effects to manifest in average noise properties.

Second, we related the dependence of HRs on target opacity to the dependence of FARs on target opacity across observers. Specifically, we correlated the z-standardized slope of the difference in HRs (Figure 2A; right panel) with the z-standardized slope of the difference in FARs (Figure 2B; right panel). This analysis yielded a moderate yet non-significant negative correlation between HR and FAR slopes (r = −.41, p = .278). In other words, observers who showed more enhancement in HRs as target opacity increased tended to show a smaller difference between congruent and incongruent FARs with increasing opacity. This correlation is based on a small number of observers (n = 9) and relates two difference measures which can be considered noisy estimates to begin with. It should thus be validated in future investigations involving a larger sample. Nonetheless, it suggests that higher target opacity levels increased observers’ sensitivity for congruent foveal orientation signals.

After establishing the impact of target opacity on HRs and FARs across all time points, we aimed to characterize the influence of target opacity on the time course of foveal enhancement. Yet, we expected target opacity to influence saccade latencies in a systematic fashion, potentially complicating a direct comparison of time courses between opacity conditions. Three interdependent factors are likely to influence the time course of congruency effects in HRs: (1) the duration for which the target had been visible when the foveal probe appeared (target-locked time course), (2) the stage of saccade preparation at which the probe was presented (saccade-locked time course) and (3) the latency of a given eye movement which moderates the relation between (1) and (2). In consequence, a probe stimulus presented, for instance, 150 ms after the target will appear in later and later stages of saccade preparation as saccade latency decreases.

Saccade parameters

Median saccade latencies decreased as the opacity and, therefore, the contrast of the saccade target increased (median mdn = 288.8, 279.7, 276.4, & 268.0 ms; Figure 3A). The slope of the linear regression line fitted to median latencies on an individual-observer level was significantly smaller than zero (mean m = −6.57 ms/opacity step, p < .001). Target opacity did not affect the precision of saccadic timing: the slope of the regression line fitted to the standard deviation of latencies on an individual-observer level was not significantly smaller than zero (m = −0.41 ms/opacity step, p = .169). To fully characterize the influence of target opacity on saccade metrics, we computed bivariate Gaussian kernel densities of observers’ saccade landing coordinates per opacity (Supplemental Figure S1B). The center of mass fell inside the target region in all conditions, demonstrating that observers were able to execute accurate saccades even at the lowest opacity. Indeed, saccade amplitudes were uninfluenced by target opacity (mdn = 10.20, 10.31, 10.25 & 10.33 dva): The slope of the regression line fitted to median amplitudes on an individual-observer level was statistically indistinguishable from zero (m = 0.03 dva/opacity step, p = .153, Supplemental Figure S1D, panel 2). Likewise, saccadic error defined as the Euclidean distance between saccade landing coordinates and the center of the target stimulus did not vary with target opacity (mdn = 1.35, 1.35, 1.38 & 1.37 dva, p = .171; Supplemental Figure S1D, panel 3). Saccadic peak velocities, on the other hand, decreased significantly with opacity (mdn = 416.7, 413.70, 410.17, 407.25 dva/s, p < .001; Supplemental Figure S1D, panel 4). However, due to the small absolute scale of these variations, the main sequence defined as the relation between saccade amplitudes and peak velocities remained largely unaltered (Supplemental Figure S1C).

Removed in the revised version: Time course of enhancement in HRs for different target opacity levels.

Earlier version: A. Probability density distributions of saccade latencies for increasing target opacity. Distributions with thin and thick lines represent individual-observer and mean probability densities, respectively. Vertical lines and shaded regions represent median latencies and SEMs, respectively. B. Target- (first row) and saccade- (second and third row) locked time course of enhancement. To obtain the corrected saccade-locked time course, the proportion of different target-locked bins in each saccade-locked bin was equalized across opacities to account for the systematic decrease in latencies with increasing opacity (see A). X-axis values indicate the center of 50 ms bins. Note that for saccade-locked time courses, the last bin contains probe onset times between 200 and 300 ms to allow for sufficient trial numbers. Across panels, error bars indicate SEMs. Asterisks denote significant differences between congruent and incongruent HRs (p≤.05; determined via bootstrapping with 10,000 repetitions; n = 9 observers).

Revised version: Time course of enhancement in HRs for different target opacity levels. A. Probability density distributions of saccade latencies for increasing target opacity. Distributions with thin and thick lines represent individual-observer and mean probability densities, respectively. Vertical lines and shaded regions represent median latencies and SEMs, respectively. B. Target-(first row) and saccade-(second and third row) locked time course of enhancement. To obtain the corrected saccade-locked time course, the proportion of different target-locked bins in each saccade-locked bin was equalized across opacities to account for the systematic decrease in latencies with increasing opacity (see A). X-axis values indicate the center of 50 ms bins. Note that for saccade-locked time courses, the last bin contains probe onset times between 200 and 300 ms to allow for sufficient trial numbers. Across panels, error bars indicate SEMs. Asterisks denote significant differences between congruent and incongruent HRs (p≤.05; determined via bootstrapping with 10,000 repetitions; n = 9 observers).

The influence of target opacity on the time course of enhancement

When probe onset was temporally aligned to the onset of the peripheral target stimulus (Figure 3B; first row), congruent and incongruent HRs did not differ significantly in any time bin for the lowest two opacities (all ps between .085 and .686; obtained with bootstrapping, see Methods). For targets with 59% opacity, congruent HRs significantly exceeded incongruent ones if the probe appeared 100-150 or 200-300 ms after the target (p = .003 & p < .001). For the highest target opacity, we observed a continuous enhancement window ranging from 100 to 200 ms after target onset (ps < .003). We subsequently inspected the saccade-locked time course by temporally aligning the offset of the probe stimulus to the onset of the eye movement. The saccade-locked development of foveal enhancement roughly mirrors target-locked time courses (Figure 3B; second row). Yet, due to the systematic shortening of saccade latencies with target opacity (see Figure 3A), the same saccade-locked time bin contained trials with systematically different target-probe intervals for different target opacities. In the last pre-saccadic bin, for instance, the probe had appeared 211 ms after the target for the lowest opacity and 178 ms after the target for the highest opacity. To control for the influence of target-probe intervals on saccade-locked time courses, we determined the proportion of each target-probe offset in every pre-saccadic time bin. We then equalized these proportions across target opacities. To achieve this, we expressed the proportions in the 59% opacity condition (i.e., approximately the opacity used in our previous investigations) as multiples of the proportions in every remaining opacity condition. We subsequently weighted the saccade-locked enhancement values by these relative proportions (see Rolfs & Carrasco, 2012 for a similar approach). As a result, enhancement in each saccade-locked bin in Figure 3C is reconstructed to contain similar probe-target offsets across opacities.

We did not observe a significant difference between congruent and incongruent HRs in any pre-saccadic time bin for the lowest target opacity (all ps between .095 and .905). At an opacity of 39%, significant enhancement manifested during the 50 ms immediately preceding saccade onset, p = .006. As target opacity increased further, enhancement was additionally observable in an earlier saccade preparation bin from 150 to 100 ms before eye movement onset, p = .003. At the highest target opacity, congruent HRs exceeded incongruent ones in the 150-100 ms bin exclusively, p = .005. In short, the corrected saccade-locked time courses appear to suggest that foveal enhancement is observable in earlier stages of saccade preparation as target opacity increases.

Earlier version: Especially at higher target opacities, the temporal development of foveal enhancement appears to exhibit an oscillatory pattern. To inspect this incidental observation in a more temporally resolved fashion, we determined mean enhancement values in a boxcar window of 50 ms duration sliding along all saccade-locked probe offset time points (step size = 10 ms; x-axis values in Figure 4 indicate the latest time point in a certain window). We then fitted 6th order polynomials (with no constraints on parameters) to the resulting time courses and determined significant time points using bootstrapping (see Methods). The average foveal enhancement across target opacities reached significance starting 115 ms before saccade onset (gray curve in Figure 4; all ps < .046). For every individual target opacity condition, we observed significant enhancement immediately before saccade onset, although only very briefly for the lowest opacity (−2–0 ms for 25%; −39–0 ms for 39%, −106–0 ms for 59% & −13–0 ms for 90%; all ps < .050; yellow to dark red curves in Figure 4). Especially for the higher two target opacities, we observed a local maximum preceding eye movement onset by approximately 80 ms. Interestingly, assuming a peak in enhancement in approximately 80 ms intervals (i.e., at x-axis values of −80 and 0 ms in Figure 4) would correspond to an oscillation frequency of 12.5 Hz. In contrast to rapid feedforward processing, feedback signaling is associated with neural oscillations in the alpha and beta range (i.e., between 7 and 30 Hz; Bastos et al., 2015; Jensen, Bonnefond, Marshall, & Tiesinga, 2015; van Kerkoerle et al., 2014).

Revised version: Especially at higher target opacities, the temporal development of foveal enhancement appears to exhibit an oscillatory pattern. To inspect this incidental observation in a more temporally resolved fashion, we determined mean enhancement values in a boxcar window of 50 ms duration sliding along all saccade-locked probe offset time points (step size = 10 ms; x-axis values in Figure 4 indicate the latest time point in a certain window). We then fitted 6th order polynomials (with no constraints on parameters) to the resulting time courses and determined significant time points using bootstrapping (see Methods). The average foveal enhancement across target opacities reached significance starting 115 ms before saccade onset (gray curve in Figure 4; all ps < .046). For every individual target opacity condition, we observed significant enhancement immediately before saccade onset, although only very briefly for the lowest opacity (−2–0 ms for 25%; −39–0 ms for 39%, −106–0 ms for 59% & −13–0 ms for 90%; all ps < .050; yellow to dark red curves in the first row of Figure 4). Especially for the higher two target opacities, we observed a local maximum preceding eye movement onset by approximately 80 ms. Interestingly, assuming a peak in enhancement in 80 ms intervals (i.e., at x-axis values of −80 and 0 ms in Figure 4) would correspond to an oscillation frequency of 12.5 Hz. In contrast to rapid feedforward processing, feedback signaling is associated with neural oscillations in the alpha and beta range (i.e., between 7 and 30 Hz; Bastos et al., 2015; Jensen, Bonnefond, Marshall, & Tiesinga, 2015; van Kerkoerle et al., 2014). We subsequently inspected the saccade-locked time course of congruent and incongruent HRs instead of their difference (Figure 4; second row). Peaks in enhancement relied on both, oscillatory increases in congruent Hit Rates and simultaneous decreases in incongruent Hit Rates. In other words, enhancement peaks appear to reflect a foveal enhancement of target-congruent feature information along with a concurrent suppression of target-incongruent features.

Removed in the revised version: Continuous time course of enhancement in HRs (HRcong–HRincong) across target opacities (gray; plot 1) and separately for different target opacity levels (plots 2–5; yellow to dark red).

Earlier version: X-axis values indicate the latest time point in a sliding boxcar window of 50 ms duration. Data points and error bars correspond to the mean ±1 SEM across observers. Lines and error bands correspond to mean 6th order polynomial fits ±1 SEM. Horizontal lines above the x-axes denote significant enhancement (p≤.05; determined via bootstrapping with 10,000 repetitions; n = 9 observers).

Revised version: First row: Continuous time course of enhancement in HRs (HRcong–HRincong) across target opacities (gray; plot 1) and separately for different target opacity levels (plots 2–5; yellow to dark red). X-axis values indicate the latest time point in a sliding boxcar window of 50 ms duration. Data points and error bars correspond to the mean ±1 SEM across observers. Lines and error bands correspond to mean 6th order polynomial fits ±1 SEM. Horizontal lines above the x-axes denote significant enhancement (p≤.05; determined via bootstrapping with 10,000 repetitions; n = 9 observers). Second row: Same plots as in the first row but separately for congruent (purple) and incongruent (gray) HRs.

The influence of saccade latency on foveal congruency effects

On an individual-trial level, the latency of the eye movement may impact foveal detection judgments beyond influencing probe timing. Since the saccade target was visible throughout saccade preparation, peripheral orientation information could accumulate for a longer period of time for long-latency saccades and may have consequently exerted a larger influence on foveal detection. On the other hand, trials with short saccade latencies likely reflect instances in which the saccade target was localized with minimal spatial uncertainty and without erroneous attentional allocation to the opposite hemifield. If short saccade latencies rely on effective target selection and oculomotor planning (see also Jonikaitis & Theeuwes, 2013; Jonikaitis, Klapetek, & Deubel, 2017; Yan, Zhaoping & Li, 2018), congruency effects may emerge most robustly on these trials. To inspect the influence of saccade latency on foveal HRs and FARs, we determined each observer’s median saccade latency per target opacity. We then separated all trials into short-latency (md = 247±18 ms) and long-latency (md = 308±18 ms) subsets, depending on whether the latency of the executed saccade on a given trial was shorter or longer than an observer’s median latency for the respective opacity level.

All previously described effects were amplified in the subset of short-latency saccades (mean n = 1,623 trials per observer): Across all time points, the difference between congruent and incongruent HRs reached significance for target opacities >=39% and increased with opacity (slope = 2.85% per opacity step, p < .001; Figure 5A). The difference in FARs, in turn, did not vary significantly with opacity (slope = 0.64% per opacity step, p = .883). Again, we observed a moderately negative yet non-significant correlation of r = −.49 (p = .181) between the normalized differences in HRs and FARs.

Earlier version: HRs and FARs separately for short-latency (A) and long-latency (B) saccades.

Revised version: HRs and FARs separately for short-latency (A) and long-latency (B) saccades. All conventions are as in Figure 2.

Removed in the revised version: All conventions are as in Figure 2.

Long-latency saccades, in contrast, showed a markedly different pattern (mean n = 1,218 trials per observer): Enhancement in HRs was overall less pronounced and even decreased with increasing target SNR (slope = −1.23% per SNR step, p = .044; Figure 5B). The differences in HRs and FARs were entirely uncorrelated (r = .09, p = .823). These findings suggest that short-latency saccades, which were enabled by effective attentional target selection and oculomotor planning, had driven the enhancement effects reported for the entire trial pool (see Figure 2). The small but significant decrease in enhancement with increasing target opacity in the long-latency subset could suggest that long-latency saccades to high-contrast targets in particular rely on attentional lapses which affected pre-saccadic perceptual processes as well as oculomotor behavior.

Discussion

We investigated whether pre-saccadic foveal congruency effects demonstrated in a previous investigation (Kroell & Rolfs, 2022) are influenced by the conspicuity of the eye movement target stimulus. For this purpose, we varied the opacity of the orientation-filtered saccade target against the unfiltered 1/f background noise. Along with the SNR of target-like orientation information, this manipulation systematically altered the luminance contrast of the target patch (see Methods). Four observations can be highlighted.

First, the manipulation of target opacity influenced saccade metrics: Saccade latencies as well as saccadic peak velocities decreased with increasing opacity. The decrease in latencies was to be expected and likely reflects a facilitation of movement planning towards high-contrast target stimuli (Ludwig, Gilchrist, & McSorley, 2004). The small yet significant decrease in peak velocities which cannot be explained by a concomitant decrease in saccade amplitudes was unexpected and, to the best of our knowledge, has not been observed previously. Hypothetically, higher peak velocities may serve to compensate for longer saccade latencies by slightly decreasing the overall time interval between target appearance and saccade landing.

Second, HRs for target-congruent and incongruent foveal probes decreased as target opacity increased, likely because attention was increasingly drawn to the target the more salient it became. Crucially, this decrease was less pronounced for congruent than for incongruent HRs, resulting in a continuous increase of enhancement with target opacity. From an experimenter’s perspective, presenting the target at a high opacity (59% and above in our design) seems beneficial since congruency effects, especially when time-resolved, are more robustly detectable. On a more conceptual level, pre-saccadic attention seems to select the target orientation in an automatic fashion, even when the eye movement could easily be programmed based on local contrast variations alone. The influence of peripheral orientation information on foveal detection judgments is directly proportional to the strength of the orientation signal at the target location. As during spontaneous eye movement behavior, saccades are routinely directed towards the currently most conspicuous object in the scene (‘t Hart et al., 2013), this finding underscores the feasibility of predictive foveal processing in natural visual environments. Of note, the relation between target opacity and foveal enhancement was exclusively driven by saccades with latencies below an observer’s median latency in each opacity condition. While the individual-trial relationship between saccade latency and the effect size of pre- and transsaccadic mechanisms is seldomly explored (but see Jonikaitis & Deubel, 2011; Jonikaitis & Theeuwes, 2013; Jonikaitis et al., 2017; Yan et al., 2018), we suggest that short saccade latencies reflect effective target selection and oculomotor planning which in turn benefits pre-saccadic perceptual processes such as foveal prediction.

Third, target opacity influenced the time course of enhancement in HRs. When systematic variations of saccade latency were taken into account (Figure 3B; third row), congruent and incongruent HRs differed earlier for higher target opacities. These findings suggest that the feedforward processing of the peripheral saccade target may have been accelerated when it was presented at high contrast (Albrecht et al., 2002). Alternatively, or in addition, it is conceivable that weaker feedforward signals require a longer accumulation interval before the feedback process can be initiated. Moreover, foveal congruency effects appear to exhibit an oscillatory pattern, with peaks in a medium saccade preparation stage (∼80 ms before the eye movement) and immediately before saccade onset. We have noticed this pattern in several investigations with substantially different visual stimuli and behavioral readouts. For instance, using a full-screen dot motion paradigm, we observed a pre-saccadic, small-gain ocular following response to coherent motion in the saccade target region (Kroell, Rolfs, & Mitchell, 2023, conference abstract; Kroell, 2023, dissertation). Predictive ocular following first reached significance ∼125 ms before the eye movement, then decreased and subsequently ramped up again ∼25 ms before saccade onset. Several explanatory mechanisms appear conceivable. Unlike rapid feedforward processing, feedback propagation has been shown to follow an oscillatory rhythm in the alpha and beta range, that is, between 7 and 30 Hz (Bastos et al., 2015; Jensen, Bonnefond, Marshall, & Tiesinga, 2015; van Kerkoerle, et al., 2014). In our case, it is possible that the object-processing areas that send feedback to retinotopic visual cortex do so at a temporal frequency of ∼12.5 Hz. At higher stimulus contrasts, feedforward signals may be fed back instantaneously and without the need for signal accumulation in feedback-generating areas. The resulting perceptual time courses may reflect innate temporal feedback properties most veridically. Alternatively, the initial enhancement peak may be related to the sudden onset of the saccade target stimulus and not to movement preparation itself. In this case, the initial peak should become particularly apparent if enhancement is aligned to the onset of the target stimulus. Yet, Figure 3 and Figure 4 suggest more prominent oscillations in saccade-locked time courses. In accordance with this, perceptual and attentional processes have been shown to exhibit oscillatory modulations that are phase-locked to action onset (e.g., Tomassini, Spinelli, Jacono, Sandini, & Morrone, 2015; Hogendoorn, 2016; Wutz, Muschter, van Koningsbruggen, Weisz, & Melcher, 2016; Benedetto & Morrone, 2017; Tomassini, Ambrogioni, Medendorp, & Maris, 2017; Benedetto, Morrone & Tomassini, 2019). Whether the oscillatory pattern of foveal enhancement, as well as its increased prominence at higher target contrasts, relies on innate temporal properties of feedback signaling, signal accumulation, saccade-locked oscillatory modulations of feedforward processing or attention, or a combination of these factors, one conclusion remains: task-induced cognitive influences suggested to underlie the considerable variability in temporal characteristics of foveal feedback during passive fixation (e.g., Fan et al., 2016; Weldon et al., 2016; 2020) are not the only possible explanation. Low-level target properties such as its luminance contrast modulate the resulting time course and should be equally considered, at least in our paradigm.

Fourth, unlike the difference between congruent and incongruent HRs, the difference in FARs did not increase with target opacity. Nonetheless, we suggest that the contribution of different response types to the pool of FAs varied across opacities: As opacity (and, therefore, the difference in HRs) increased, incidental visual properties in the foveal background noise more reliably reflected the subsequently reported orientation. In other words, FAs were increasingly triggered by signal, that is, orientation information presented on screen, as target opacity increased. Moreover, observers who showed more enhancement in HRs as target opacity increased tended to show a smaller difference between congruent and incongruent FARs with increasing opacity. Combined with a range of arguments provided in our previous publication (Kroell & Rolfs, 2022), this observation suggests that foveal congruency effects reflect a variation of pre-saccadic foveal sensitivity rather than a shift in a post-perceptual decision criterion. As noted in the Results section, however, this finding is based on a small number of observations and should be confirmed in larger samples.

Foveal prediction in natural visual environments

As noted above, human observers typically move their eyes towards the most conspicuous objects in their environment. Foveal prediction seems to benefit from this strategy as the strength of the predicted signal increases with the conspicuity of the eye movement target. Nonetheless, natural visual environments as well as natural viewing behavior pose several challenges for the foveal prediction mechanism (see Kroell & Rolfs, 2022, for an initial discussion).

First, naturalistic saccade target stimuli will likely exhibit complex shapes and, more often than not, will include feature conjunctions rather than isolated features. Previous findings suggest that the foveal feedback mechanism is capable of operating at this level of complexity: High-level peripheral information such as the category of novel, rendered objects has been successfully decoded from activation in foveal retinotopic cortex (e.g., Williams et al., 2008). If, indeed, temporal object-specific areas such as area TE send feedback, the foveal prediction mechanism may even be specialized for the transfer of complex visual properties.

Second, foveal input will often be of high contrast in natural visual environments. If fed-back predictive signals can influence foveal perception in the presence of high-contrast feedforward input remains to be established. In our main investigation (Kroell & Rolfs, 2022; Figure 2B) as well as in previous studies (Hanning & Deubel, 2022b), pre-saccadic foveal detection performances decreased markedly in the course of saccade preparation, presumably because visuospatial attention gradually shifted towards the saccade target and away from the foveal location. This pre-saccadic decrease in foveal sensitivity may boost the relative weight of fed-back signals by attenuating the conspicuity of high-contrast feedforward input. In other words, the strength of feedforward input to the fovea is reduced gradually across saccade preparation. At the same time, the strength of the fed-back predictive signal should profit from the high contrast of naturalistic saccade targets.

Third, while foveal and peripheral information was congruent on 50% of all ‘probe present’ trials in our investigation, peripheral and foveal features will often be weakly correlated (see Samonds, Geisler, & Priebe, 2018) or even uncorrelated in natural environments. Again, the pre-saccadic attenuation of foveal feedforward processing may allow fed-back peripheral signals to influence perception even if they are uncorrelated with foveal information. Moreover, in piloting variations of our paradigm, we observed that the subjective impression of perceiving the saccade target at the pre-saccadic foveal location is most pronounced if the foveal noise region is replaced with a black Gaussian blob at certain time points before saccade onset (observation by author LMK; unpublished). In consequence, fed-back signals do not seem to require correlated feedforward input to influence perception. Quantitative evidence, however, remains to be established.

Lastly, pre-saccadic foveal input is likely less relevant during natural viewing behavior than it is in our task. It is possible that this task-induced prioritization of the foveal location facilitated the emergence of congruency effects. In a previous experiment (Kroell & Rolfs, 2022; Figure 1D), however, the perceptual probe could appear anywhere on a horizontal axis of 9 dva length around the screen center. Despite this spatial unpredictability, however, congruency effects peaked at the pre-saccadic foveal location, even after peripheral baseline performances had been raised to a foveal level through an adaptive increase in probe opacity. Ultimately, an influence of task demands on visual processing can only be fully excluded through techniques that provide a direct readout of perceptual contents without requiring keyboard responses. In psychophysical investigations, a prediction of saccade target motion may be read out from observers’ eye velocities (Kroell, Mitchell, & Rolfs, 2023; Kwon, Rolfs, & Mitchell, 2019). In electroencephalographic (EEG) and neurophysiological studies, foveal predictions should manifest in early visually evoked potentials (e.g., Creel, 2019) and increased firing rates of feature-selective foveal neurons in early visual areas, respectively.

Neural implementation of foveal prediction

Based on the body of our findings as well as previous literature, we suggested a parsimonious feedback mechanism to underly the observed effects: the preparation of a saccadic eye movement, and the concomitant shift of pre-saccadic attention (e.g., Kowler, Anderson, Dosher, & Blaser, 1995; Deubel & Schneider, 1996), selects the peripheral target stimulus among competing information. Higher-order visual areas feed selected feature input back to early retinotopic areas—specifically, to neurons with foveal receptive fields. Fed-back feature information combines with congruent, foveal feedforward input, resulting in the enhancement effects we observe. Especially in the context of active vision, this feedback mechanism is appealing as it resolves a combinatorial issue associated with feature-specific information transfer before saccades. Consider a simplified case in which, right before a saccadic eye movement, the activation of a feature-selective neuron that encodes a certain retinal location is transferred to a neuron within the same brain area that will encode said retinal location after saccade landing. For this mechanism to function for any possible saccade direction and amplitude, most neurons would need to be connected to most other neurons (or, in a simplified version, to neurons with foveal receptive fields) in a given brain area. Assuming an information transmission via feedback rather than horizontal connections significantly reduces this dimensionality: Higher-order visual areas that encode object properties (largely) detached from retinotopic or spatiotopic reference frames selectively transfer feature information to neurons with foveal receptive fields, irrespective of the vector of the upcoming saccade. This parsimonious mechanism would have shortcomings. In particular, foveal feedback should become less effective during saccade sequences where several peripheral targets are simultaneously attended. Feature information at both attended target locations may be fed back in temporal succession or weighted and erroneously combined into a single fed-back signal. In most cases, however, foveal feedback may reasonably achieve what established transsaccadic mechanisms struggle to explain: An anticipation of the features of a single saccade target—which typically constitutes the currently most relevant object in the visual field—in foveal vision.

While direct feedback connections from higher-order to early visual areas would constitute the most straightforward implementation, it is conceivable that feedback signals are relayed through and modulated by subcortical areas. In particular, the thalamic pulvinar has been identified as a connection hub for visual processing that receives copies of feedforward and feedback connections from different visual areas and may even combine information across visual space (Cortes, Ladret, Abbas-Farishta, & Casanova, 2024). In the case of foveal prediction, thalamic neurons may receive fed-back signals from higher-order areas and enhance those signals before passing them on to cortical neurons with foveal receptive fields. Perhaps, a modification of foveal activation within the thalamic pulvinar itself is sufficient to influence perception. To the best of our understanding, however, the fed-back signal must originate in non-retinotopic, higher-order object processing areas to reduce the number of necessary neuronal connections.

Conclusion

To sum up, this investigation both provides further mechanistic insights into the pre-saccadic foveal prediction mechanism and constrains the parameter space for researchers planning to adapt our paradigm in future studies. Conceptually, active peripheral feature sampling is not necessary for foveal prediction to emerge. Instead, foveal congruency effects develop even when (or especially when) salient local contrast variations at the saccade target location can be used to direct the eye movement. In consequence, presenting the target at a high opacity appears purely beneficial. At high target opacities, foveal congruency effects in HRs are more pronounced. Moreover, observers respond more systematically to target-like orientations in foveally presented noise, facilitating reverse-correlation based analyses.

Methods

Sample

Nine human observers (six females, eight right-handed, six right-eye dominant, no authors) aged 22 to 30 years (mdn = 26.0) participated in the experiment. Normal (n = 3) or corrected-to normal (n = 6) visual acuity was ensured at the beginning of the first session using a Snellen chart (Hetherington, 1954) embedded in a Polatest vision testing instrument (Zeiss, Oberkochen, Germany). Observers yielding scores of 20/25 or 20/20 were invited to proceed with the experiment. Ocular dominance was assessed using the Miles test (Miles, 1930). Since data collection was performed during the COVID-19 pandemic, our sample was composed of lab members. Nonetheless, all observers were naïve as to the purpose of the study. Participants gave written informed consent before the experiments and were compensated with either accreditation of work hours or a payment of 8.50€/hour plus a bonus of 1€/session. The study complied with the Declaration of Helsinki in its latest version and was approved by the Ethics Committee of the Department of Psychology at Humboldt-Universität zu Berlin. The research question, experimental paradigm and data analyses were preregistered on the Open Science Framework (link will be made available upon acceptance). Raw and pre-processed data as well as the experimental code are available in the same repository.

Experimental setup and design

The external setup, session structure and trial time course were identical to our previous experiments (Kroell & Rolfs, 2022) with the following exceptions: First and most importantly, the saccade target was overlayed on the background noise with one of four logarithmically spaced opacities (α = 25%, 39%, 59% and 90% compared to 60% in Kroell & Rolfs, 2022). Targets with different opacities were presented in separate experimental blocks the order of which was randomized. Before every block, we showed a preview of the target stimulus at the respective opacity (sequentially at both possible target locations in the left and right screen half, randomly oriented either 45° to the left or right). Second, during each trial in a block, we removed the target upon saccade initiation, i.e., once gaze had crossed a boundary of 2.0 dva diameter around the fixation dot. Third, we increased the number of potential background noise images from 17 to 27 images per session and observer. All 27 images were generated at the beginning of a given session. Before each trial, we randomly drew 16 images from this pool (without replacement) and presented them in random order. Since 9 observers each completed 7 sessions, 27*9*7 = 1701 different noise images were presented in the course of the experiment. Lastly, observers completed four staircase blocks before the main experiment in every session (one block per target opacity). Individual staircase blocks were short (32 trials each) and served to familiarize observers with all target opacities as well as to obtain an estimate of observers’ typical saccade latency for each opacity. Saccade latency estimates were used to define three possible probe onset times for each target opacity in the main experiment. Importantly, however, the same staircase ran throughout all four blocks and returned one optimal probe opacity estimate that was used throughout the main experiment, irrespective of target opacity.

Removed in the revised version: Variation of RMS contrast (first column), Michelson contrast (second column) and SNR (third column) within the saccade target region.

Removed in the revised version: A. Probability density distributions per measure and target opacity (yellow to dark red shadings). B. Mean and standard deviation of contrast and SNR separately for each noise image presented during the saccade preparation period. The target was presented from the 4th noise image on.

Earlier version: Increasing target opacity increased the luminance contrast and the SNR of the filtered orientation within the saccade target region. During the experiment, we manipulated target opacity by varying the ‘globalAlpha’ input to the Psychtoolbox function ‘DrawTexture’ (see http://psychtoolbox.org/docs/Screen-DrawTexture) and enabled alpha blending with the inputs ‘GL_SRC_ALPHA’ (source) and ‘GL_ONE_MINUS_SRC_ALPHA’ (destination; see http://psychtoolbox.org/docs/Screen-BlendFunction). To emulate the functionality of this presentation technique (see https://github.com/Psychtoolbox-3/Psychtoolbox-3/blob/master/Psychtoolbox/ PsychDemos/ContrastModulatedNoiseTheElegantStyleDemo.m; all links retrieved on 2/9/23), we determined the luminance sum of the background noise and the orientation-filtered target patch multiplied with the respective opacity: where ‘luminanceCombined’ denotes the resulting luminance values within the target region, ‘luminanceBackground’ denotes the luminance values within a square of 3 dva width around the center of the target location, ‘opacity’ denotes the manipulated target opacity (25-90%), ‘luminanceFiltered’ denotes the luminance values of the orientation-filtered target patch (3 dva width) and ‘cosineMask’ denotes the raised cosine mask with which the filtered patch was overlayed (3 dva diameter). We then calculated the root mean square (RMS, see Peli, 1990) and Michelson contrast (Michelson, 1927) of the resulting luminance sum for every trial. To compute the SNR of the target orientation, we determined the energy of various SF and orientation combinations in the combined luminance patch (see Data analysis). We then divided the mean filter energy within a region of +/-15° around the target orientation by the mean filter energy for all remaining orientations. An increase in target opacity from 25% to 90% corresponded to an increase in RMS contrast from .08 to .11, an increase in Michelson contrast from .53 to .78, and an increase in the SNR of the filtered orientation from 1.40 to 2.58 (Figure 5).

Revised version: Increasing target opacity increased the luminance contrast and the SNR of the filtered orientation within the saccade target region. During the experiment, we manipulated target opacity by varying the ‘globalAlpha’ input to the Psychtoolbox function ‘DrawTexture’ (see http://psychtoolbox.org/docs/Screen-DrawTexture) and enabled alpha blending with the inputs ‘GL_SRC_ALPHA’ (source) and ‘GL_ONE_MINUS_SRC_ALPHA’ (destination; see http://psychtoolbox.org/docs/Screen-BlendFunction). To emulate the functionality of this presentation technique (see https://github.com/Psychtoolbox-3/Psychtoolbox-3/blob/master/Psychtoolbox/PsychDemos/ContrastModulatedNoiseTheElegantStyleDemo.m; all links retrieved on 2/9/23), we determined the luminance sum of the background noise and the orientation-filtered target patch multiplied with the respective opacity: where ‘luminanceCombined’ denotes the resulting luminance values within the target region, ‘luminanceBackground’ denotes the luminance values within a square of 3 dva width around the center of the target location, ‘opacity’ denotes the manipulated target opacity (25-90%), ‘luminanceFiltered’ denotes the luminance values of the orientation-filtered target patch (3 dva width) and ‘cosineMask’ denotes the raised cosine mask with which the filtered patch was overlayed (3 dva diameter). We then calculated the root mean square (RMS, see Peli, 1990) and Michelson contrast (Michelson, 1927) of the resulting luminance sum for every trial. To compute the SNR of the target orientation, we determined the energy of various SF and orientation combinations in the combined luminance patch (see Data analysis). We then divided the mean filter energy within a region of +/-15° around the target orientation by the mean filter energy for all remaining orientations. An increase in target opacity from 25% to 90% corresponded to an increase in RMS contrast from .08 to .11, an increase in Michelson contrast from .53 to .78, and an increase in the SNR of the filtered orientation from 1.40 to 2.58 (Figure 6).

Added in the revised version: Variation of RMS contrast (first column), Michelson contrast (second column) and SNR (third column) within the saccade target region. A. Probability density distributions per measure and target opacity (yellow to dark red shadings). B. Mean and standard deviation of contrast and SNR separately for each noise image presented during the saccade preparation period. The target was presented from the 4th noise image on.

Data analysis

Online trial abortion criteria as well as trial exclusions based on offline gaze analyses are identical to our previous study (Kroell & Rolfs, 2022). Unlike in our previous investigation, we intended to remove the saccade target during saccadic flight on all trials. Across observers, 0.24% (min = 0%; max = 1.1%; std = 0.40%) of trials were excluded due to the saccade target still being visible after saccade landing.

Earlier version: To test the influence of target opacity on the difference between congruent and incongruent HRs and FARs, we performed linear mixed-effects models in which we described the variance in both differences (HRcong–HRincong ; FARcong–FARincong) across all time points with a fixed effect of target opacity and a random intercept for observer: section, this model outperformed the simplest model including a fixed effect of target opacity only (1) and a more complex one involving a random intercept and random slope for observer (2):

Revised version: To test the influence of target opacity on the difference between congruent and incongruent HRs and FARs, we performed linear mixed-effects models in which we described the variance in both differences (HRcong–HRincong ; FARcong–FARincong) across all time points with a fixed effect of target opacity and a random intercept for observer: where ‘cong-incong’ refers to the difference between either congruent and incongruent HRs or congruent and incongruent FARs. As stated in the Results section, this model outperformed the simplest model including a fixed effect of target opacity only (1) and a more complex one involving a random intercept and random slope for observer (2):

Model fitting was performed with the Matlab function fitlme (Matlab 2020b, Mathworks, Natick, MA, USA).

All remaining tests of statistical significance relied on bootstrapping: Within each observer, we determined the means in the to-be compared conditions and computed the difference between those means. Across observers, we drew 10,000 random samples from these differences (with replacement). Reported p-values correspond to the proportion of differences smaller than or equal to zero. We considered p-values ≤ 0.05 significant.

Furthermore, to relate the difference between congruent and incongruent HRs to the difference in FARs, we z-standardized HR and FAR slopes before computing the Pearson correlation coefficient.

Earlier version: Lastly, we used a three-step approach to calculate the reverse correlation index plotted in Figure 2C (see Figure 6). First, we determined all noise images that had been visible during saccade preparation on trials in which observers generated a congruent or incongruent FA. We then described the average visual properties of these noise images within a 3 dva diameter circular region around the screen center using a set of Gabor filters with varying SF*orientation properties [SFs from .67 to 2.0 cycles per degree (cpd) in 20 equal steps of .07 cpd; orientations from −90° to 90° in 13 equal steps of 15°; see Movellan, 2002; Wyart, Nobre, & Summerfield, 2012; Li, Barbot, & Carrasco, 2016; Schweitzer and Rolfs, 2020, 2021]. Second, we subtracted the filter energies on incongruent FA trials from filter energies on congruent FA trials on an individual-observer level. Third, we selected a 30° orientation window around the non-target orientation (i.e., −45°) and subtracted the mean filter energy in that window from the mean filter energy in a 30° orientation window around the target orientation (i.e., 45°). The higher the resulting value, the more congruent and incongruent FAs were based on signal, i.e., on incidental orientation information on screen. Note that we combined the lower two and the higher two target opacities to increase the number of trials and obtain reliable reverse-correlation measures. The resulting index is based on m = 2,461 (std = 583) and m = 1,952 (std = 629) trials per observer for the lower two and higher two opacities, respectively.

Revised version: Lastly, we used a three-step approach to calculate the reverse correlation index plotted in Figure 2C (see Figure 7). First, we determined all noise images that had been visible during saccade preparation on trials in which observers generated a congruent or incongruent FA. We then described the average visual properties of these noise images within a 3 dva diameter circular region around the screen center using a set of Gabor filters with varying SF*orientation properties [SFs from .67 to 2.0 cycles per degree (cpd) in 20 equal steps of .07 cpd; orientations from −90° to 90° in 13 equal steps of 15°; see Movellan, 2002; Wyart, Nobre, & Summerfield, 2012; Li, Barbot, & Carrasco, 2016; Schweitzer and Rolfs, 2020, 2021; details provided in Kroell & Rolfs, 2022]. Second, we subtracted the filter energies on incongruent FA trials from filter energies on congruent FA trials on an individual-observer level. Third, we selected a 30° orientation window around the non-target orientation (i.e., −45°) and subtracted the mean filter energy in that window from the mean filter energy in a 30° orientation window around the target orientation (i.e., 45°). The higher the resulting value, the more congruent and incongruent FAs were based on signal, i.e., on incidental orientation information on screen. Note that we combined the lower two and the higher two target opacities to increase the number of trials and obtain reliable reverse-correlation measures. The resulting index is based on m = 2,461 (std = 583) and m = 1,952 (std = 629) trials per observer for the lower two and higher two opacities, respectively.

Removed in the revised version: Calculation of the reverse correlation index plotted in Figure 2C, illustrated for the 59% target opacity condition.

Earlier version: In step 1, we identified the average properties (SF*orientation) of the foveal noise window on congruent (purple outlines and font) and incongruent (gray outlines and font) FA trials. In step 2, we determined the difference between them (congruent–incongruent). In step 3, we subtracted filter energies around the non-target orientation (−45°) from filter energies around the target orientation (45°).

Revised version: Calculation of the reverse correlation index plotted in Figure 2C, illustrated for the 59% target opacity condition. In step 1, we identified the average properties (SF*orientation) of the foveal noise window on congruent (purple outlines and font) and incongruent (gray outlines and font) FA trials. In step 2, we determined the difference between them (congruent–incongruent). In step 3, we subtracted filter energies around the non-target orientation (−45°) from filter energies around the target orientation (45°).